```python
# Initialize Otter
import otter
grader = otter.Notebook("hw1.ipynb")
```

# Hw 1: Data Science Workflow and More NumPy 📈

Name:

Student ID:

### Collaborators and Sources
In addition to recording your **collaborators** (TAs, peers, group members, roommates, etc.) on this homework in the cell below, you are required to **cite/indicate all external sources** used when finishing this assignment. 
External sources are defined as anything that is not considered course material, such as online sources (webpages, blog posts, etc.), content generated by AI systems, books, etc.

When using external sources, indicate what kind of external sources (e.g. stack overflow, WashU chatGPT, etc.) you used in the cell below and then provide more specific citations (such as links to webpages or in case of AI generated asnwers the actual sources (links provided by the AI system) and the prompt used to generate the answers) with your answer to **each specific problem**.  


Note that these citations will not free you from your obligation to submit your _own_ code and write-ups, however, they will be taken into account during the grading and regrading process, **especially** when two or more submissions closely resemble each other. Working with each other is ok, as long as you cite who you worked with and you don't copy anyone's answers directly!

Collaborators:
* Albert Einstein
* Marie Curie

External Sources: 
* google search (then provide specific links in respective problem such as: https://developers.google.com/edu/python/strings)

### Submission instructions
* Submit this Python notebook on Gradescope.
* Execute all cells and save your notebook prior to submission. 
* Don't forget to commit and push to your Github repo!
* **Do not change the number of cells!** Your submission notebook should have exactly one code cell per problem. 
* Do **not** remove the `# your code here` line. Add your solution **after** that line. 

## Let's get started!

For this homework, we will be gaining more exposure to NumPy and exploring the data science workflow. We will also practice plotting data.  **Important:** Ensure that all plots have the necessary components that a plot should have: axes labels and a title (and a legend when it is appropriate). You must have all components of a plot on every plot for the entire semester!

Work your way through these problems. They will also provide some more practice for working with NumPy arrays. 
> Use NumPy arrays and functions in **all** problems unless explicitly stated otherwise! 

Frequently **save** your notebook!


### Learning Goals
* **[G1-1]** Be able to list the **steps of the DS Workflow**. Know (and be able to describe) what happens in each step.
* **[G1-4]** Apply the data science and machine learning workflow to a **real-world problem**. 
* **[G1-5]** **Represent data** mathematically and programmatically.
* **[G1-6]** Analyze and visualize **continuous features**.
* [G0-6] Understand that arrays are multi-dimensional data structures of fixed dimensions. _Think of arrays as vectors._ 
#### Acquire Python skills
* [G0-5] practice syntax, data structures, control-flow
* work with `numpy` arrays
* get exposed to `matplotlib`

## 1. Churn Analysis: Data Science Workflow

Consider the following data science use case:

> *Assume you just landed a great analytical job with MegaTelCo, one of the largest telecommunication firms in the United States. They are having a major problem with customer retention in their wireless business. In the mid-Atlantic region, 20% of cell phone customers leave when their contracts expire, and it is getting increasingly difficult to acquire new customers. Since the cell phone market is now saturated, the huge growth in the wireless market has tapered off. Communications companies are now engaged in battles to attract each other’s customers while retaining their own. Customers switching from one company to another is called churn, and it is expensive all around: one company must spend on incentives to attract a customer while another company loses revenue when the customer departs.*
>
> *You have been called in to help understand the problem and to devise a solution. Attracting new customers is much more expensive than retaining existing ones, so a good deal of marketing budget is allocated to prevent churn. Marketing has already designed a special retention offer. Your task is to devise a precise, step-by-step plan for how the data science team should use MegaTelCo’s vast data resources to decide which customers should be offered the special retention deal prior to the expiration of their contracts.*

<!-- BEGIN QUESTION -->


### Problem 1.1

**Write-up!** Is this a scientific, social, or business problem? Briefly justify your answer.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Problem 1.2

**Write-up!** Using the **data science workflow** discussed in class, describe how you would takle this task. Make sure to list every step and briefly describe the step _in the context of this specific use case_. 

> **Hint**: Include ideas on what data you might use and how they would be used. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 2. Practice NumPy Indexing

Indexing, especially with NumPy, can be a tricky feature to truly wrap one's head around, but the benefits of (working towards) mastering it make it a worthy endeavor. The more practice you get, the easier it will become — eventually you won't be able to even imagine how else you would do things.

In this section, we will start with some review and then move on to more complex features.

### Basic Indexing
At the risk of belaboring this topic, let's quickly do some practicing. The following cell produces `yet_another_array` using the `np.random.randint` function, which we will use in the next few problems.

```python
import numpy as np
yet_another_array = np.random.randint(11, size=10)
yet_another_array
```

### Problem 2.1

**Do this!** Retrieve the 4th element of `yet_another_array`.

*Note: this question has hidden tests, or is graded on style of code and not just answer alone.*

```python
element = ...
```

### Problem 2.2

Here's something new: specifically for NumPy arrays, you can also pass in a Python list or NumPy array of indicies to retrieve.

For example, `some_array[[0, 2, 4, 6]]`.


**Do this!** Retrieve the 2nd, 5th, and 9th elements of `yet_another_array`.

*Note: this question has hidden tests, or is graded on style of code and not just answer alone.*

```python
elements = ...
```

```python
grader.check("q2b")
```

### Problem 2.3

Let's do this with a 2D array, too. Here we will generate `a_2d_array` using the `np.random.rand` function.

*Note: this question has hidden tests, or is graded on style of code and not just answer alone.*

```python
a_2d_array = np.random.rand(5, 5)
a_2d_array
```

Remember that you can index into a 2D array like this: `another_2D_array[row, column]`, where `row` and `column` are indicies, slices, `:`s, or some mix of these three.

**Do this!** Retrieve the value at index `(3, 4)` from `a_2d_array`.

```python
a_2d_element = ...
```

### Slice Indexing

You've already gotten familiar with slices, but here is some more practice.

### Problem 2.4

**Do this!** Retrieve the 6th, 7th, and 8th elements of `yet_another_array`.

*Note: this question has hidden tests, or is graded on style of code and not just answer alone.*

```python
more_elements = ...
```

```python
grader.check("q2d")
```

### Problem 2.5

Here is a reminder of what `a_2d_array` looks like.

*Note: this question has hidden tests, or is graded on style of code and not just answer alone.*

```python
a_2d_array
```

**Do this!** Retrieve the values in the 3rd column of `a_2d_array`.

```python
column = ...
```

```python
grader.check("q2e")
```

### Boolean/Logical Indexing

Now for something complete new! You might be interested in finding all of the values in an array that fulfills some condition like _get all values greater than 5_. 

> [🐍 **Python Feature** 🐍]: NumPy enables this by supporting [**Boolean (or logical) indexing**](https://docs.scipy.org/doc/numpy-1.13.0/user/basics.indexing.html#boolean-or-mask-index-arrays).

This idea is also referred to as **masking**. You can use an array of boolean values _of the same shape_ as the target array as the "index" into the target array. This will return all of the values that are in the same position as the `True` values from the boolean arrays. These logical arrays are called "masks" because they are analogous to masks, which let some of the underlying surface show through, but hide the rest.

Let's try making our first logical array. In the cell below, we use a comparison operation on `yet_another_array`, returning an array of the same size with `True` in the positions where the values meet the condition.

```python
yet_another_array > 5
```

Now let's see how this array "lines up" with `yet_another_array`.

```python
yet_another_array
```

Here is another example using an array that is not sorted already.

```python
shuffled_array = np.random.permutation(yet_another_array)
shuffled_array
```

```python
shuffled_array > 5
```

Is this what you would expect? Let's actually get those values.

```python
shuffled_array[shuffled_array > 5]
```

### Problem 2.6

**Do this!** Retrieve the elements of `yet_another_array` that are even (**Hint**: `%`).

*Note: this question has hidden tests, or is graded on style of code and not just answer alone.*

```python
even_array = ...
```

### Boolean/Logical Indexing for Machine Learning
Let's assume we have some data for supervised machine learning stored in `X` (_features_) and `y` (_targets_). 

```python
y = np.array([0, 2, 2, 2, 1, 0, 0, 1, 2, 0, 1, 0, 2, 1, 2, 1, 1,
                   0, 2, 0, 0, 1, 0, 2, 1, 1, 0, 1, 1, 0, 1, 2, 2, 0, 
                   0, 0, 2, 1, 0, 0, 2, 2, 0, 1, 1])
np.random.seed(21)
X = np.random.rand(len(y),4)
```

### Problem 2.7

**Do this!** Assume you have your features `X` and a categorical target `y`. Retrieve all fearures for all datapoints that have a target value of `2`. Store them in `features_of_class_2`.

*Note: this question has hidden tests, or is graded on style of code and not just answer alone.*

```python
features_of_class_2 = ...
```

```python
grader.check("q2g")
```

## 3. Model Evaluation and More Numpy


### Error Measure

Lets assume we have the following predictions `y_pred` and observed targets `y_true` for evaluation. 

```python
y_hat = np.array([0, 2, 1, 2, 1, 0, 0, 2, 1, 0, 1, 0, 2, 1, 2, 1, 1, 
                   0, 2, 0, 0, 1, 0, 2, 1, 1, 0, 1, 1, 0, 1, 2, 2, 0, 
                   0, 0, 2, 1, 0, 0, 2, 2, 0, 1, 1])
y_test = np.array([0, 2, 2, 2, 1, 0, 0, 1, 2, 0, 1, 0, 2, 1, 2, 1, 1,
                   0, 2, 0, 0, 1, 0, 2, 1, 1, 0, 1, 1, 0, 1, 2, 2, 0, 
                   0, 0, 2, 1, 0, 0, 2, 2, 0, 1, 1])
```

### Problem 3.1

**Do this!** Complete the following function so that given two arrays of labels, `y_test` and `y_hat`, it will return the average classification error, `avg_err`. Use **Boolean operations** on the arrays instead of a for loop.

*Note: this question has hidden tests, or is graded on style of code and not just answer alone.*

```python
def avg_error(y_true, y_pred):
    '''
    Computes the classification error given Y_TRUE and Y_PRED
    as a value from 0 to 1
    '''
    
    avg_err = ...
    
    return avg_err
```

```python
err_score = avg_error(y_hat, y_test)
print(f'Error: {100*err_score:0.2f}%')
```

<!-- BEGIN QUESTION -->

<div class="alert alert-block alert-warning">
<b>You are required to complete Lab1</b> for the following problem(s). If your group was not able to complete Lab1 within Tuesday's lab time. Complete it now.
</div>

**Write-up!** How does the _average classification error_ computed above relate to the _average classification accuracy_ that you computed in Lab 1? State a mathematical formula of how you can compute the average error from the average accuracy. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Evaluating a Model using all the Features

As you might have concluded in your communication your findings step in Lab1, one way to improve our predictions could be to use all the features instead of just one. Let's do that now and compare the results. We saved the processed data for you. This is the same data that you created and used in Lab1. 

```python
data = np.load('utility/data/processed_iris_data.npy')
data[:5,:]
```

### Problem 3.2

**Do this!** Create `X` and `y` arrays. Now with X including all four features. Then reproduce the steps from Lab1 (D. Building and evaluating a model and making predictions). 

*Note: this question has hidden tests, or is graded on style of code and not just answer alone.*

```python
X, y = "Your", "Answer"
...
```

### Problem 3.3

**Do this!** Make the training and testing sets exactly with the same configuration as in Lab1. 
> Note: This is **very important** since it will let us compare the new results with the ones we got in the lab (since we are using the _same_ training and test splits).

*Note: this question has hidden tests, or is graded on style of code and not just answer alone.*

```python
from sklearn.model_selection import train_test_split

# train_test_split will be used a lot in this class, be sure to remember how to use it!
X_train, X_test, y_train, y_test = "Your", "Answer", "Goes", "Here"
...
```

### Problem 3.4

**Do this!** Train the Model on the training data that you created in the previous step. 

*Note: this question has hidden tests, or is graded on style of code and not just answer alone.*

```python
from utility.util import SimpleModel
model = ...
```

```python
...
```

```python
grader.check("q3d")
```

### Problem 3.5

**Do this!** Make predictions on the test set and compare them to the ground truth by computing the average classification error. 
> **Hint:** use your `avg_error()` function from above. 

*Note: this question has hidden tests, or is graded on style of code and not just answer alone.*

```python
predictions = ...

score = ...

print(f'Avg Error: {100*score:0.2f}%')
```

```python
grader.check("q3e")
```

<!-- BEGIN QUESTION -->

### Problem 3.6

**Write-up!** Report your average classification error from both the lab and this homework. Did the the model in this homework perform better or worse compared to the lab? Why do you think the performace imporved or did not imporve?    

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**That's it! Follow the submission instructions below:**
* Restart your kernel and run all the cells
* Save your notebook
* Upload the .ipynb notebook directly to Gradescope

```python

```

Tests for q2a (points: 5)


Tests for q2b (points: 5)
```
>>> len(elements) == 3
True
```


Tests for q2c (points: 5)


Tests for q2d (points: 5)
```
>>> len(more_elements) == 3
True
```


Tests for q2e (points: 5)
```
>>> len(column) == 5
True
```


Tests for q2f (points: 4)


Tests for q2g (points: 8)
```
>>> np.isclose(features_of_class_2[0, 0], X[y == 2][0, 0]) == True
np.True_
```
```
>>> np.all(features_of_class_2.shape == (13, 4)) == True
np.True_
```


Tests for q3ai (points: 6)


Tests for q3b (points: 6)


Tests for q3c (points: 6)


Tests for q3d (points: 6)
```
>>> from sklearn.neighbors import KNeighborsClassifier
>>> model._clf.get_params(False)['n_neighbors'] == 1
True
```


Tests for q3e (points: 10)
```
>>> len(predictions) == 45
True
```
