```python
# Initialize Otter
import otter
grader = otter.Notebook("hw2.ipynb")
```

# Hw 2: More Pandas and Data Exploration! 🐼

Name:

Student ID:

### Collaborators and Sources
In addition to recording your **collaborators** (TAs, peers, group members, roommates, etc.) on this homework in the cell below, you are required to **cite/indicate all external sources** used when finishing this assignment. 
External sources are defined as anything that is not considered course material, such as online sources (webpages, blog posts, etc.), content generated by AI systems, books, etc.

When using external sources, indicate what kind of external sources (e.g. stack overflow, WashU chatGPT, etc.) you used in the cell below and then provide more specific citations (such as links to webpages or in case of AI generated asnwers the actual sources (links provided by the AI system) and the prompt used to generate the answers) with your answer to **each specific problem**.  


Note that these citations will not free you from your obligation to submit your _own_ code and write-ups, however, they will be taken into account during the grading and regrading process, **especially** when two or more submissions closely resemble each other. Working with each other is ok, as long as you cite who you worked with and you don't copy anyone's answers directly!

Collaborators:
* Albert Einstein
* Marie Curie

External Sources: 
* google search (then provide specific links in respective problem such as: https://developers.google.com/edu/python/strings)

### Submission instructions
* Submit this Python notebook on Gradescope.
* Execute all cells and save your notebook prior to submission. 
* Don't forget to commit and push to your Github repo!
* **Do not change the number of cells!** Your submission notebook should have exactly one code cell per problem. 
* Do **not** remove the `# your code here` line. Add your solution **after** that line. 

## Let's get started

For this homework, work through **Lab 2 (Analyzing the MoMA Data)** first. In general, you should feel free to import any package that we have previously used in class. Ensure that all plots have the necessary components that a plot should have (e.g. axes labels, a title, and a legend if it is applicable).

Frequently **save** your notebook!

### Learning Goals
* **[G1-6]** Analyze and visualize continuous features.
* **[G2-5]** Implement and follow the EDA Process. 
* **[G2-7]** Practice posing and answering data questions.
* **[G2-8]** Know common functions to compute summary statistics from data using pandas. 
#### Acquire Python skills
* use `pandas` DataFrames 
* plotting with `matplotlib`
* save & load data with `pickle`, file handling

### Some imports and configurations

```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
```

```python
# Let's run some configurations to make prettier plots
from utility.util import configure_plots

configure_plots()
```

```python
%%javascript
IPython.OutputArea.prototype._should_scroll = function(lines) {
    return false;
}
```

## 1. More MoMA Data Visualizations

In `Lab2` you visulaized _categorical_ features in the MoMA data. Let's look into visualizing _continuous_ features.

Let's read the data first. 

<div class="alert alert-block alert-warning">
<b>You are required to complete Lab2</b> for the following problem(s). If your group was not able to complete Lab2 within Tuesday's lab time. Complete it now.
</div>

[*ungraded*] **Do this!** Go to your Lab2 folder, run the Lab2 notebook. Using a the _terminal_ or a _file browser_ (yes, I mean the file browser of your operating sysyem, **not** the jupyter file hierarchy view in the browser window) move the files `processed_artworks.pickle` and `MoMADirectorsDepartmentHeads.csv` from `lab2-XYZ/utility/data` to `hw2-ABC/utility/data`. Then run the cell below to load the data. 

> **Hint**: `mv source_path_to_filename destiation_path_to_filename` using the command line. 

**DO NOT EDIT THE FOLLOWING CELL, THE AUTOGRADER DEPENDS ON THIS PATH!**

```python
artworks = pd.read_pickle("./utility/data/processed_artworks.pickle")

directors = directors = pd.read_csv("./utility/data/MoMADirectorsDepartmentHeads.csv")
```

<!-- BEGIN QUESTION -->

### Problem 1.1

Let's investiage the **size of the MoMA artworks**, specifically their **widths**.

**Write-up!** What type of variable is the _width of the artwork_ and which plot is best used to visualize this feature.  

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Problem 1.2

**Do this!** Compute and print the following summary statistics of the artwork width feature: _mean, standard deviation, range (min, max)_. 

> **Best practice**: When filtering NaN values (e.g., using `notnull()`) or removing outliers (see below) it's best practice to only perform these on the features you are currently looking at. Do **not** remove whole data entries from the dataframe.

*Note: this question has hidden tests, or is graded on style of code and not just answer alone.*

```python
# Extract the width column
w = artworks['Width_cm']

# Do not consider examples where the value = NaN
id_NaN = artworks['Width_cm'].notnull()
w = w[id_NaN]

mean = ...
standard_dev = ...
max_width = ...
min_width = ...

print(f'mean: {mean}')
print(f'standard deviation: {standard_dev}')
print(f'min: {min_width}')
print(f'max: {max_width}')
```

<!-- BEGIN QUESTION -->

### A Question to Investigate 

Something looks interesting here: What's up with artworks having a width of 9000cm (90 meters = 295 ft)? 

* **(Q1)** How many artworks have large widths, let's say a width of more than 8000cm? 


Let's approach this with a plot.

### Problem 1.3

**Do this!** Create an appropriate plot to visualize the frequency of artworks with respect to their width. Select an appropriate non-default _binwidth_.  
> **Reminder**: do not forget to include all necesary components in your plot: _title, axis labels, units_ (if applicable). This will hold for all plots you will create in this and future assignemnts. We will not add this reminder in the future. 

```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Write-up!** This write-up has two parts: 
* Describe your plot: What does it show? 
* Critically question your work: Is this a good visualization? Does it help in answering **(Q1)**?  Why/why not? 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Problem 1.4 (a)

Let's change the scale of the plot.

**Do this!** By using an appropriate `?` command to retrieve the documentation of your plot function used above, find out how to create the same plot using a **log scale** (for the y axis). Create the plot leaving all other parameters the same as in your first plot. Assign it an appropriate title (that is differnt from the plot above). 

> **IMPORTANT:** comment out the line using the `?` command to avoid long print outs while executing your notebook prior to submission. 

```python
...
```

<!-- END QUESTION -->

### Problem 1.4 (b)

<!-- BEGIN QUESTION -->

**Write-up!** Answer **(Q1)**: How many artworks have large width, let's say a width of more than 8000cm? 
> **Hint**: You do not have to run any code to find this answer. Simply look at the plot above.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Problem 1.5

Instead of the plot using the log-scale, we can also take another approach: we can remove the _outliers_.

**Do this!** Create a temporary variable `w_new` of the width of all artworks with width smaller than 250cm. Also, a width of 0 doesn't really make much sense either; remove artworks with width 0cm! Visualize the frequency of artworks with respect to their width not including these outliers.  


> **Hint**: Use **Boolean indexing** on `w`! This approach is required to receive full credit.

```python
...
```

<!-- END QUESTION -->

### Problem 1.6
Hold on!!! Now our summary statistics have changed. 

**Do this!** Compute the summary statistics from this new variable: _mean, standard deviation, and range (min, max)_. 

*Note: this question has hidden tests, or is graded on style of code and not just answer alone.*

```python
new_mean = ...
new_standard_dev = ...
new_max_width = ...
new_min_width = ...

print("New Data stats: ")
print(f'mean: {new_mean}')
print(f'standard deviation: {new_standard_dev}')
print(f'min: {new_min_width}')
print(f'max: {new_max_width}')
```

<!-- BEGIN QUESTION -->

### Problem 1.7

**Write-up!** Look at all plots again, briefly describe the _shape of the distributions_ (use the **technical terms** to efficiently and  precisely summarize this infomration). 
> **Hint**: This is covered in zyBook section 3.3.
 
Which of the three approaches/data visualizations do you think is best? Briefly justify your choice by including one benefit of the approach you selected and one disadvantage of each of the two approaches you did not deem appropriate/best. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Good to Know: Many Distributions have a _Long Tail_

In statistics and business, a _long tail_ of some distributions of numbers is the portion of the distribution having many occurrences far from the "head" or central part of the distribution. The distribution could involve popularities, random numbers of occurrences of events with various probabilities, etc. The term is often used loosely, with no definition or arbitrary definition, but precise definitions are possible. Read more on [Wikipedia](https://en.wikipedia.org/wiki/Long_tail)!

#### Why does the long tail matter? 

Read this story on [WIRED](https://www.wired.com/2004/10/tail/).

## 2. Investiage Hypothesis (Q2) from Lab2

Recall, (Q2) was _How does the gender distribution of acquired works change with different museum directors?_ 

First, we'll have to do some data cleaning on the `directors` table, such as extracting the MoMA directors. Then we will add the `Director_ID` to the `artworks` table. 

```python
# Extract MoMA directors
moma_directors = directors[directors.DepartmentFullName=='The Museum of Modern Art']
moma_directors
```

```python
# Add director ID of director acquiring the work to artworks table
def get_director(year):
    """Given the year of artwork acquisition provide the director ID."""
    for director_ID in moma_directors.index:
        if year < moma_directors.sort_values(
                by="PositionBeginYear").loc[director_ID, "PositionEndYear"]:
            return int(director_ID)
    return int(director_ID)

try:
    artworks = pd.read_csv("./utility/data/artworks_with_director.csv", )
    artworks['DateAcquired'] = pd.to_datetime(artworks['DateAcquired'])
    print("Loaded precomputed data.")
except:
    print("This will take a couple of minutes.")
    artworks["Director_ID"] = artworks.DateAcquired.apply(lambda x: get_director(x.year))
    artworks.to_csv("./utility/data/artworks_with_director.csv", index=False)
    artworks['DateAcquired'] = pd.to_datetime(artworks['DateAcquired'])
    print("Completed cell block")
```

Now, let's compute the fraction of artworks by female artists acquired by each MoMA director normalized by the length of each director's appointment (using `groupby("Director_ID")` and `mean()`). Based on these numbers we can then get the fractions of artworks from female artists per year. 

```python
frac_female_by_dir = artworks.groupby("Director_ID").frac_females.mean()
new = moma_directors.join(frac_female_by_dir)
new[["DisplayName","frac_females"]]
```

<!-- BEGIN QUESTION -->

### Problem 2.1

**Do this!** Create a bar plot of MoMA directors versus the fraction of acquired artworks by female artists. 

> Hint: check the example visualizing a categorical feature we provided in **Step 3: Data Profiling** of Lab2.

```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Problem 2.2

**Write-up!** Summarize your results on the female gender representation by differnt museum directors. Highlight, anything you find notable or surprising and keep a note on any further questions you might have. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 3. World Happiness: Understanding the Data

It's time to explore another dataset. This time the data comes from _socioeconomics_ and attempts to decribe the state of global happiness. The dataset we will be using is from the [World Happiness Report from 2023](https://worldhappiness.report/ed/2023/).

<img src="utility/pics/WorldHappiness.png" alt="Drawing" style="width: 500px;"/>

In the following section, we will do some preliminary exploration. First, find the data download link in your repo's `README.md`. Store the `happiness.csv` data in `utility/data`. Then, let's load the data into `happiness`.

**DO NOT EDIT THE FOLLOWING CELL, THE AUTOGRADER DEPENDS ON THIS PATH!**

```python
import os

data = 'utility/data'

assert os.path.exists(data), \
    'Unable to find data; make sure that you have the `utility` directory'

happiness = pd.read_csv(f'{data}/happiness.csv')
```

<!-- BEGIN QUESTION -->

### Problem 3.1 (a)

Let's get to know our data set.

**Do this!** Use the cell below to investigate the data set as you answer the following questions.

```python
...
```

**Write-up!** How many **observations** are in our dataset? How many **features** does each one have?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Problem 3.1 (b)

<!-- BEGIN QUESTION -->

**Write-up!** What are the **features**? What **types of variables** (categorical, continuous, etc.) are each of them? How are they represented programmatically (i.e., which **data type** – string, float, etc. – do you use to represent each feature)?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Problem 3.2

**Write-up!** What does this data set seem to be about? If this data set was used to build a model, which feature could be the **target**? Give a brief **justification** for your selection of the target variable. Which would be the **input variables**? 

Here is the data set again.

```python
happiness.head()
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 4. Exploring Individual Happiness Features

Next, let's take a look at the features individually.

### Problem 4.1

**Do this!** In the cell below, create an indexer that can be used to get the numerical columns of the `happiness` DataFrame and store it in `numerical`. For example, if we evaluated `happiness[numerical].head()`, we would get all columns with names in `numerical`.

*Note: this question has hidden tests, or is graded on style of code and not just answer alone.*

```python
...
numerical = ...

happiness[numerical].head()
```

<!-- BEGIN QUESTION -->

### Problem 4.2 (a)

**Do this!** Use the cell below to investigate the data set as you answer the following questions.

```python
...
```

**Write-up!** What are the average values for each numeric feature? Yes, list them all.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Problem 4.2 (b)

<!-- BEGIN QUESTION -->

Recall, standard deviation is given by
$$\text{std}(x) = \sigma = \sqrt{\dfrac{\sum_{i=1}^n (x_i - \mu)^2}{n-1}}.$$

**Write-up!** Which feature has the largest standard deviation? Which one has the lowest?

> Note that this is just for practicing how to retrieve these values. Since the scales for each may be different, we cannot meaningfully compare each standard deviation value.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Problem 4.3

Let's implement a median finding function from scratch.

**Do this!** Complete the function below so that it returns the `median` of an array `x`. You may use NumPy indexing operations to do this, but **do not use NumPy's built-in median function**.

*Note: this question has hidden tests, or is graded on style of code and not just answer alone.*

```python
def compute_median(x, axis=0):
    '''
    Returns the median for an array X.
    AXIS specifies the axis of the array across which the median is copmuted.
    '''
    median = ...
    ...
    
    return median
```

```python
grader.check("q4cii")
```

<!-- END QUESTION -->

Let's try using your implementation to compute the median values for each feature.

**Do this!** Use the [`apply`🔗](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.apply.html) method on DataFrames to compute the medians for each feature and store the result in `medians`.

*Note: this question has hidden tests, or is graded on style of code and not just answer alone.*

```python
medians = ...

medians
```

```python
grader.check("q4ciii")
```

<!-- BEGIN QUESTION -->

### Problem 4.4

Let's create some visualizations.

**Do this!** In the cell below, create box plots for the following 4 features: `Social Support`, `Freedom to make life choices`, `Generosity`, `Perceptions of corruption`. You can use the `.values` property of DataFrames to get a NumPy array. Ensure that your plots have all of the necessary components.

> **[Optional]**: Try using a [`violin plot`🔗](https://matplotlib.org/api/_as_gen/matplotlib.pyplot.violinplot.html?highlight=violinplot#matplotlib.pyplot.violinplot) to visualize each feature. You can use this function the same way as you would use `boxplot`.

```python
...
```

**Write-up!** Using this visualization, describe the _shape of each of the feature distributions_. Do you notice any issues with this visualization if your goal is to _compare_ the distributions? 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Problem 4.5

Let's look at one of those features in more detail. 

**Do this!** In the cell below, create a histogram plot with 15 bins of the feature `Perceptions of Corruption`.

```python
...
```

**Write-up!** Describe the distribution including its _shape_ and what it means in the context of the _specific data_ it shows. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 5. Comparing Pairs of Happiness Features

Now that we have a sense of what the features look like, we can start examining their relationships with each other.

### Problem 5.1

**Do this!** Complete the function below so that it will display a scatter plot, given a data frame and two column names. Makre sure you have all plot components. 
> **Hint**: Check `Lab1` for some sample code if you need a reminder.)

```python
def scatter_columns(df, first_column, second_column):
    '''
    Creates a scatter plot of two columns of a data frame.
    '''
    ...
```

<!-- BEGIN QUESTION -->

**Do this!** Using the function `scatter_columns` from above, explore **various combinations** of features to see which ones show some kind of trend or pattern. You must analyze more than 2 combinations! If your chosen combinations do not show much trend, analyze more to find one with noticeable trend!

```python
...
```

**Write-up!** Let's look at the pairwise feature comparison. 
* Which pairs of features show the most trend (makre sure to display it above)? Describe the trend. 
* Given this information, what can you conclude? 
* **[ungraded]** This is the time to make note of anything surprising or interesting you found when looking at these plots. This way you can remember later and investiage this further or use this information to guide your subsequient steps or "connect the dots" when interpreting/discussing results. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Pearson Correlation

There are several ways to quantify trends, but let's pick one to discuss today. The [Pearson Correlation Coefficient](https://en.wikipedia.org/wiki/Pearson_correlation_coefficient) (commonly denoted by $\rho$ or $r$) is a frequently used statistic for quantifying linear relationships between two variables (or features in our case). It has a value of $+1$ if there is a perfect positive linear trend, a value of $-1$ if there is a perfect negative linear trend, and a value of $0$ if there is no trend.

<img src="utility/pics/correlation.png" alt="correlation" width="80%"/>

We can compute the sample correlation coefficient, $r$, with the following formula:
$$
r(X, Y) = \dfrac{\text{Cov}(X, Y)}{\sigma_X \sigma_Y} = \dfrac{\frac{1}{n-1} \sum_{i=1}^n (x_i - \mu_X)(y_i - \mu_Y)}{\sigma_X \sigma_Y},
$$
where $X$ and $Y$ are variables, $\mu$ is a mean, and $s$ is a sample standard deviation,
$$
\sigma_X = \sqrt{\dfrac{\sum_{i=1}^n (x_i - \mu_X)^2}{n-1}}.
$$

<!-- BEGIN QUESTION -->

### Problem 5.2

**Do this!** Complete the `correlation` function in the following cell so that it computes the correlation coefficient between two arrays `x` and `y`. Remember to use **vectorized operations**, but for the sake of this excercise **do not use numpy's built in functions for variance, covariance, standard deviation, or the correlation coefficient**. Using the numpy `mean` function is allowed. 
> **Hint**: It may be helpful to work out the math for the equations. You can start by ["foiling"](https://en.wikipedia.org/wiki/FOIL_method) the numerator.

*Note: this question has hidden tests, or is graded on style of code and not just answer alone.*

```python
def correlation(x, y):
    '''
    Returns the sample Pearson Correlation Coefficient between arrays X and Y.
    '''
    std = ...
    corr = ...
    ...
    
    return corr

correlation(happiness.Score, happiness.Generosity)
```

```python
grader.check("q5bii")
```

<!-- END QUESTION -->

Let's try using your implementation to compute some correlation coefficients.

**Do this!** In the cell below, compute the correlation coefficient values between `Score` and every other feature, including score itself, and store them in `corrs`. 
> **Hint**: Since this is an operation pairwise in the number of variables (which are typically considered a small number) rather than in the number of data points (which we typically consider a rather large number), it is totally ok to use `for` loop here. 

> **Hint**: Use `numerical` from [Problem 4.1](#Problem-4.1) and the `.columns` property of DataFrames to get the names of the columns that you need.

*Note: this question has hidden tests, or is graded on style of code and not just answer alone.*

```python
corrs = ...

corrs
```

Let's visualize these correlation values. To do this, we will use [Seaborn's](https://seaborn.pydata.org) `heatmap` function.

```python
import seaborn as sns

corrs = np.array(corrs)
plt.figure(figsize=(1, 6))
sns.heatmap(corrs[:,None], cmap=sns.cm.vlag, vmin=-1)
plt.xticks([0.5], ['Score'], rotation=0)
plt.yticks(np.arange(corrs.shape[0]) + 0.5, happiness.loc[:,numerical].columns, rotation=0)
plt.title('Correlation: Score vs Rest')
```

As you may have seen in the test at the end of the cell before the previous cell, we can get the _Correlation Matrix_ of pairwise correlation values by using the `.corr` method on DataFrames.

```python
happiness.corr(numeric_only=True)
```

```python
sns.heatmap(happiness.corr(numeric_only=True), cmap=sns.cm.vlag, vmin=-1)
plt.xticks(rotation=45)
plt.title('Happiness: Correlation Matrix')
```

<!-- BEGIN QUESTION -->

**Write-up!** Answer the following summary questions: 
* Which features (apart form `Score` itself) have the **highest** correlation with Score? 
* Which have the **lowest**? 
* Are there **negative correlations**?
* Do you think these results are surprising? Briefly justify why/why not?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**That's it! Follow the submission instructions below:**
* Restart your kernel and run all the cells
* Save your notebook
* Upload the .ipynb notebook directly to Gradescope

```python

```

Tests for q1b (points: 3)


Tests for q1f (points: 3)


Tests for q4a (points: 5)


Tests for q4cii (points: 5)
```
>>> x = np.random.rand(10)
>>> np.isclose(compute_median(x), np.median(x))
np.True_
```
```
>>> x = np.random.rand(9)
>>> np.isclose(compute_median(x), np.median(x))
np.True_
```


Tests for q4ciii (points: 2)
```
>>> np.all(medians == happiness.median(numeric_only=True))
np.True_
```


Tests for q5bii (points: 5)
```
>>> np.isclose(correlation(happiness.Score, happiness.Generosity), np.corrcoef(happiness.Score, happiness.Generosity)[0, 1])
np.True_
```


Tests for q5biii (points: 3)
