```python
# Initialize Otter
import otter
grader = otter.Notebook("hw3.ipynb")
```

# Hw 3: More Linear Regression 🎢

Name:

Student ID:

### Collaborators and Sources
In addition to recording your **collaborators** (TAs, peers, group members, roommates, etc.) on this homework in the cell below, you are required to **cite/indicate all external sources** used when finishing this assignment. 
External sources are defined as anything that is not considered course material, such as online sources (webpages, blog posts, etc.), content generated by AI systems, books, etc.

When using external sources, indicate what kind of external sources (e.g. stack overflow, WashU chatGPT, etc.) you used in the cell below and then provide more specific citations (such as links to webpages or in case of AI generated asnwers the actual sources (links provided by the AI system) and the prompt used to generate the answers) with your answer to **each specific problem**.  


Note that these citations will not free you from your obligation to submit your _own_ code and write-ups, however, they will be taken into account during the grading and regrading process, **especially** when two or more submissions closely resemble each other. Working with each other is ok, as long as you cite who you worked with and you don't copy anyone's answers directly!

Collaborators:
* Albert Einstein
* Marie Curie

External Sources: 
* google search (then provide specific links in respective problem such as: https://developers.google.com/edu/python/strings)

### Submission instructions
* Submit this Python notebook on Gradescope.
* Execute all cells and save your notebook prior to submission. 
* Don't forget to commit and push to your Github repo!
* **Do not change the number of cells!** Your submission notebook should have exactly one code cell per problem. 
* Do **not** remove the `# your code here` line. Add your solution **after** that line. 

## Let's get started

For this homework, work through **Lab 3 (Predicting Housing Prices with Linear Regression)** first. Most of the things we ask you to do in this homework are explained in the lab. In general, you should feel free to import any package that we have previously used in class. Ensure that all plots have the necessary components that a plot should have (e.g. axes labels, a title, and a legend when applicable).

Frequently **save** your notebook!

### Learning Goals
* **[G3-6]** Practice **model evaluation and comparison**. Be able to compute RMSE and MAE and use it for model evaluation.
* **[G3-7]** Use `sklearn` to create a **train-test split**, **train a multiple linear regression model**, and **make predictions**. 
* **[G1-4]** Apply the data science and machine learning workflow to a real-world problem. 
* **[G2-4]** Practice data exploration, understanding, and visualization. 
#### Acquire Python skills
* use `sklearn` for ML

## 0. Reviewing Regression

The following questions are ungraded. We leave them in here for you to answer them to solidify your understanding of M3 learning goals (aka linear regression) and as prep for the learning goal quiz. 

**Q1:** What distinquishes regression tasks from classification and clustering tasks? Consider the input features and outputs of models from each type.

**Q2:** Describe the procedure for "fitting" a linear regression model. Assuming that we have $d$ predictors, how many parameters should be fit for linear regression? How do we find the optimal parameters for fitting our training data?

**Q3:** Given two linear models, how do we decide which one fits the data _better_?

**Q4:** What is training error? What is testing error? What are their interpretations in the context of the data and what do they mean (in the context of the data)?

**Q5:** Consider two models: the first is a linear model; the second is a polynomial model. The second model is more complex. Describe a situation where the polynomial model performs better (has less error) during training but the linear model performs better during testing. What is the name for what happened to the polynomial model?

## 1. Linear Regression with More Features

In this week's lab, we implemented a linear regression model that could map a single predictor feature to a target variable. This is also called _simple linear regression_. In the case of the Boston housing dataset, we could map features like the "average number of rooms per dwelling" to a house's price. In reality, however, a single predictor is not enough to get a good model. Additionally, we have data for several predictors so why not make use of that?

In this series of problems, let's explore how we can build a multi-dimensional linear model also called _multiple linear regresison_.

### Feature Matrix X and Target Vector y

Let's load the data and create our feature _matrix_ and target _vector_.

> **Note**: A _matrix_ is a 2-dimensional array. In **mathematical terms** we speak of _scalars_, _vectors_, and _matrices_ rather than single/multiple valued variables or 1D and 2D arrays. Also note that we typically use the following notation convention:
* $x$, $y$, $a$, $b$, ... denote _scalars_, 
* **x, y, a, b** or $\vec{x}$, $\vec{y}$, $\vec{a}$, $\vec{b}$, ... denote _vectors_, and 
* $X$, $Y$, $A$, $B$, ...  denote _matrices_. 

> We even follow this convention for variable names: lowercse for scalar- and vector-valued variables (e.g., `x,y,a,b`) and uppercase for matrices/2D arrays, e.g., `X`. 

```python
from utility.util import load_boston
boston = load_boston()
```

```python
X, y = boston.data, boston.target
```

### Scikit-Learn

Generalizing our derivation of a 1D linear regression to a multi-dimensionional one is non-trivial since we would need to convert our computations into matrix operations, which are beyond the scope of this class. However, that does not mean that we can't do multi-dimensional linear regressions. The package we will introduce today is [Scikit-learn](https://scikit-learn.org/stable/). It contains many generalized implementations of common learning models and we will be using its implementation of multi-dimensional linear regression here.


Using `sklearn` is very similar to our implementation from the lab, though there is a slight difference: we work with a model object. In the following cell, we create an instance of linear regression model called `model`.

```python
from sklearn.linear_model import LinearRegression

model = LinearRegression()
```

The `fit` and `predict` functions that we used previously are analogous to the `fit` and `predict` methods of `model`. Here is a quick example of how to perform 1D linear regression, mapping the `CRIM` predictor to housing price.

```python
feature_to_use = 'CRIM'
x_crim = X[:, boston.feature_names.tolist().index(feature_to_use)].reshape(-1, 1)


# you should create a new model object for each new model you create
model = LinearRegression()

# fitting the model
model.fit(x_crim, y)

# predicting with the model
import numpy as np

x_star = np.linspace(x_crim.min(), x_crim.max(), 1001).reshape(-1, 1)
model.predict(x_star)
```

<!-- BEGIN QUESTION -->

### Problem 1.1

**Write-up!** Recall the **phases of the supervised ML workflow** and how data is used in each of them. What is missing in the example code provided above? Elaborate on both data and phases. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Problem 1.2

Now that we are familiar with `sklearn`, let's start building a linear model that uses all of the acceptable features in our dataset. You can refer to the [documentation](https://scikit-learn.org/stable/modules/generated/sklearn.linear_model.LinearRegression.html) for more information about how to use `LinearRegression`.

**A Note on Language:** The Boston dataset used to be provided by the sklearn library (now hosted on Kaggle) with its original data description including the feature description. You will notice that this description uses insensitive language. We do not endorse this terminology. The use of this dataset, however, is intentional to raise awareness of ethical considerations and to sharpen our minds to spot ethical issues while making decisions using data. The broader purpose of the exercise is to recognize that any use of the data that encodes race, gender, or other such categories, even if phrased in the most respectful way, has the potential to introduce harmful bias into data-driven decision making and machine learning models.

Recall our discussion on dataset bias from Lab 3 and how the use of certain data/features has the potential to add biases to our machine learning models. For a more thorough discussion, check out the analysis [here](https://fairlearn.org/main/user_guide/datasets/boston_housing_data.html).

**Do this!** Create the _feature matrix_ `X_fair2use`, which includes all data points and all features specified as `True` in a **Boolean mask vector** `features_to_use`. Exclude the follwing features:

- `'B'` recording the proportion of black people by town
- `'LSTAT'` recording the percentage of "lower status" people in the population
    
The reason we want to do this is to remove possible bias/predicting "power" to be based on race or socioeconomic status. 

> Hint: There is a section on _Masking: Using Boolean Arrays to create Mask Arrays_ in the **Numpy cheatsheet**. 

*Note: this question has hidden tests, or is graded on style of code and not just answer alone.*

```python
features_to_use = ...
X_fair2use = ...
```

```python
grader.check("q1bi")
```

**Do this!** Build a model that uses _all_ of the remaining predictor features in the Boston housing dataset. Remember to create a new instance of `LinearRegression` and assign it to the `model` variable. Also make sure that you do a train/test split of your data so we can evaluate our model after training (use `random_state=10` so we can compare our results to the ones from `Lab3`; refer to the lab for an example of how to do this). If you were successful, you will see the weights for the model below.

*Note: this question has hidden tests, or is graded on style of code and not just answer alone.*

```python
from sklearn.model_selection import train_test_split


X_train, X_test, y_train, y_test = ...

model = ...

...

# Get the model parameters
w = np.append(np.array(model.intercept_), model.coef_)
print(w)

# Now, w is a vector! With the bias term plus a weight parameter for every feature.
print(w.shape)
```

```python
grader.check("q1bii")
```

### Problem 1.3

Now that we have a trained model, we can use it to make some predictions using the testing set.

**Do this!** Use your model to compute the predicted prices of the points in your testing set and store these into the variable `y_prediction`. Then, compute the RMSE of these predictions and store the value in `rmse`. If all went well, you should see the RMSE value you computed as the output of the cell. 

> **Hint**: You can use the sklearn mean_squared_error function

*Note: this question has hidden tests, or is graded on style of code and not just answer alone.*

```python
...

y_prediction = ...

rmse = ...

f'RMSE: {rmse}'
```

```python
grader.check("q1c")
```

<!-- BEGIN QUESTION -->

### Problem 1.4

Let's talk about this result.

**Write-up!** Summarize your findings by describing your new result and how it compares to those of individual predictors from `Lab3`? Is this what you expected to see and why?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 2. Introducing a New Evaluation Metric ($R^2$) and Discussion

RMSE is a good measure for how well a model fits, but one could argue it is not the most intuitive. An alternative measure is $R^2$, or the **coefficient of determination**, which is essentially a measure of how close the data are to the fitted regression line. RMSE is a number in the units of the original target variable. $R^2$ is a value between 0 and 1, where increasing values indicate better fit (1 means a perfect fit). In this way, $R^2$ is a good place to start when evaluating a regression model because it is easy to interpret.

$$R^2 = 1 - \frac{\displaystyle \frac{1}{n} \sum_{n=1}^n (y_i - f(x_i))^2}{\displaystyle \frac{1}{n} \sum_{n=1}^n (y_i - \bar y)^2}$$

> _For those who are interested_: Those of you that are statistically savvy, or read the lab very thoroughly (which you should do regardless!!), will remember that these sums are mean squared errors. The numerator is the MSE of the true values compared to the predicted value and the denominator is the MSE of the true value relative to the mean value. The denominator is the variance of the response variable (e.g. house price). One way to interpret this metric is that it computes the ratio between how much variation is explained by the model and the total variation in the data.

### Problem 2.1

**Do this!** Complete the function below so that it computes and returns the $R^2$ value of a model given some `predictions` and their true `labels`. To receive full credit, implement the math yourself; do not rely on other, already implemented functions. Store your computed result in the `r2` variable. Then, use the function to compute the $R^2$ value of your model from Problem 2.2 on the testing data, and save the $R^2$ value as `score`.

> Do **not** use any built-in functions from scikit-learn for this problem. 

```python
def r_squared(predictions, labels):
    '''computes the r-squared metric of a model given some PREDICTIONS and their true LABELS'''
    
    r2 = ...
    
    return r2

score = ...
```

```python
grader.check("q2a")
```

<!-- BEGIN QUESTION -->

### Problem 2.2

Now, we have a good prediction model! Do we? Let's be critical and second guess the use of the Boston Housing data to build a predictive model for house prices. Read the following description of the data: https://fairlearn.org/main/user_guide/datasets/boston_housing_data.html.

**Answer the following questions in your _own words_. No GenAI!** 

**Write-up!** When and why was the data collected and what was its intentional use?  

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Write-up!** What are the issues with this dataset? Name and briefly descirbe at least two issues. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Write-up!** Why is this dataset not suitable for building a prediciton model even if we omitted the problematic features? 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Write-up!** How could you build a viable predictor to predict housing prices in boston?  

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 3. Using a More Complex Model

There are a lot of options we can choose from when it comes to picking a model to use when analyzing our data. In this section, we will take a look at a more complex model: **polynomial regression**. We will also investigate some of the considerations one should make when picking a model.

![poly_reg](utility/pics/poly-regression.png)

### Back to Toy Data

As mentioned in earlier, it is really hard to visualize multi-dimensional datasets, so for this section, we will fall back to the toy data we used from `Lab3`.

```python
import numpy as np
x, y = np.loadtxt('utility/data/toy_data.txt')
```

Polynomial regression models are really just extensions of the linear models that we are already familiar with. The difference comes from how we preprocess our data. The specifics of how this works is beyond the scope of this course (see CSE417T for more information or ask a TA).

In order to build a polynomial model, we will need to do some preparation. `sklearn`'s models require that the input data is an $n$-dimensional array. Since our input array `x` is a 1D array for the toy data set, we will need to reshape it into a 2D array. The lambda function provided in the following cell, `reshape`, will take care of this for you — use it when applicable (see Lab 0 for a refresher on lambda functions if needed).

```python
reshape = lambda x: x.reshape(-1, 1)
```

<!-- BEGIN QUESTION -->

### Problem 3.1

The next thing we will need to do is fit a linear model that will serve as our baseline.

**Do this!** Create a new linear model called `linear_model` and fit it with the toy data set and create a plot showing the scattered data points `x, y` and the predicted values of the model (don't forget all of the necessary plot components!). You don't need to worry about making a train/test split for this section.

```python
import matplotlib.pyplot as plt
from utility.util import configure_plots

configure_plots() # for pretty plots

...
```

<!-- END QUESTION -->

### Doing Polynomial Regression

Next, we will provide a function that produces fitted polynomial regression models. Again, you don't need to worry too much about how this works, but you are welcome to try to understand it.

```python
from sklearn.pipeline import Pipeline
from sklearn.preprocessing import PolynomialFeatures

def make_poly(x, y, degree=2):
    '''creates and returns a polynomial regression model fit with input data X and Y'''
    
    polynomial_features = PolynomialFeatures(degree=degree)
    poly_model = Pipeline([("polynomial_features", polynomial_features),
                           ("linear_regression", LinearRegression())])

    return poly_model.fit(reshape(x), y)
```

<!-- BEGIN QUESTION -->

### Problem 3.2

Let's try using the `make_poly` function.

**Do this!** Use `make_poly` to, well, make a polynomial regression model with a `degree` of 2. This model will fit the data with a quadratic function. Store the model you created in `poly2_model` and create a plot showing the scattered data points `x, y` and the predicted values of the model (with all of the necessary plot components!). 
> **Hint**: you will need to make an array of equally spaced values over an interval (see [`np.linspace`](https://docs.scipy.org/doc/numpy/reference/generated/numpy.linspace.html)) for your predicted values; refer to `Lab3`.

```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Problem 3.3

Nice, that worked well! Let's up the complexity again. Try to play around with the degree and find the one that produces the "best-fitting" model.

**Do this!** Use `make_poly` and try to find the "best-fitting" model (based on $R^2$) by experimenting with the `degree` argument. Store this model in `poly_best_model`. Then, create a plot showing the scattered data points `x, y` and the predicted values of the model. Make sure your graph has the proper components. 

> **Hint**: use the `r_squared` metric you implemented before to evaluate the fit of your model.

```python
...
```

<!-- END QUESTION -->

### Evaluating the Models

Let's compare the RMSE and $R^2$ metrics for the models that we have produced.

*Note: this question has hidden tests, or is graded on style of code and not just answer alone.*

### Problem 3.4

**Do this!** Fill in the following `for` loop to print out the RMSE and $R^2$ metrics for each model in `models`.

```python
models = {
    'Linear': linear_model,
    '2nd Order Polynomial': poly2_model,
    'Best Fitting Polynomial': poly_best_model
}

model_eval = {}
RMSE_IDX = 0
R2_IDX = 1

for name, toy_model in models.items():
    
    RMSE = ...
    R2 = ...
    
    model_eval[name] = (RMSE, R2)
    
    
# now we print our values
for name, evals in model_eval.items():
    print(name)
    print(f'RMSE: {evals[RMSE_IDX]:0.3f}')
    print(f'R2: {evals[R2_IDX]:0.3f}')
    print()
```

<!-- BEGIN QUESTION -->

### Problem 3.5

**Write-up!** Describe the restuls. What do you notice about the scores? Given these results, which model would you pick to deliver to a client? Why?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Stacking Up Against the Real Function

What if I told you the true function was $$y = 0.16x + 120$$

The following function, `ground_truth`, will return `x` and `y` values produced by the true function with some added noise. We can use this to produce other points that could have been generated by the same function.

```python
def ground_truth(n):
    '''returns N random samples from the function above'''
    
    w, b = 0.16, 120
    noise = np.random.rand(n) * 350
    noise = noise - np.mean(noise)
    x = np.random.randint(0, 2000, n)
    y = w * x + b + noise
    
    return x, y
```

<!-- BEGIN QUESTION -->

### Problem 3.6

**Do this!** With `ground_truth`, make a plot that shows these five things (and all plot components! We will stop reminding you about including these from here on out, so be sure to not forget them for all plots you create moving forward!):
1. The points from the toy dataset
2. 200 points generated from the true function
3. The linear model
4. The "best-fit" polynomial model
5. The true function

You don't need to retrain the models for this part. To receive full points, ensure that your plot has all of the proper components, including a legend (see `Lab3` for an example of how to do this).

```python
x_truth, y_truth = ground_truth(200)
```

```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Problem 3.7

**Write-up!** Describe the results plotted in [Problem 3.6](#Problem-3.6). How well do both models fit the training vs. the test data?  What is the danger of fitting a model on a limited number of training examples? 

Conclude by briefly summarizing which model is better and why in one sentence.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**That's it! Follow the submission instructions below:**
* Restart your kernel and run all the cells
* Save your notebook
* Upload the .ipynb notebook directly to Gradescope

```python

```

Tests for q1bi (points: 5)
```
>>> features_to_use
array([ True,  True,  True,  True,  True,  True,  True,  True,  True,
        True,  True, False, False])
```
```
>>> X_fair2use.shape == (506, 11)
True
```
```
>>> len(features_to_use) == 13
True
```


Tests for q1bii (points: 5)
```
>>> w.shape == (12,)
True
```


Tests for q1c (points: 5)
```
>>> np.isscalar(rmse) == True
True
```


Tests for q2a (points: 6)
```
>>> np.isscalar(r_squared(y_prediction, y_test)) == True
True
```
```
>>> np.isclose(r_squared(y_prediction, y_test), model.score(X_test, y_test)) == True
np.True_
```
```
>>> np.isclose(score, model.score(X_test, y_test)) == True
np.True_
```


Tests for q3d (points: 6)
