```python
# Initialize Otter
import otter
grader = otter.Notebook("hw0.ipynb")
```

# Hw0: What is DS and Intro To Python 🐍

Name:

Student ID:

### Collaborators and Sources
In addition to recording your **collaborators** (TAs, peers, group members, roommates, etc.) on this homework in the cell below, you are required to **cite/indicate all external sources** used when finishing this assignment. 
External sources are defined as anything that is not considered course material, such as online sources (webpages, blog posts, etc.), content generated by AI systems, books, etc.

When using external sources, indicate what kind of external sources (e.g. stack overflow, WashU chatGPT, etc.) you used in the cell below and then provide more specific citations (such as links to webpages or in case of AI generated asnwers the actual sources (links provided by the AI system) and the prompt used to generate the answers) with your answer to **each specific problem**.  


Note that these citations will not free you from your obligation to submit your _own_ code and write-ups, however, they will be taken into account during the grading and regrading process, **especially** when two or more submissions closely resemble each other. Working with each other is ok, as long as you cite who you worked with and you don't copy anyone's answers directly!

Collaborators:
* Albert Einstein
* Marie Curie

External Sources: 
* google search (then provide specific links in respective problem such as: https://developers.google.com/edu/python/strings)

### Submission instructions
* Submit this Python notebook on Gradescope.
* Execute all cells and save your notebook prior to submission. 
* Don't forget to commit and push to your Github repo!
* **Do not change the number of cells!** Your submission notebook should have exactly one code cell per problem. 
* Do **not** remove the `# your code here` line. Add your solution **after** that line. 

## Let's get started!

For this homework, work through **Lab 0 (Introduction to Python)** first. Most of the things we ask you to do in this homework are explained in the lab. In general, you should feel free to import any package(s) that we have previously used in class. **IMPORTANT:** Ensure that all plots have the necessary components that a plot should have: axes labels and a title (and a legend when it is appropriate). You must have all components of a plot on every plot for the entire semester!

Frequently **save** your notebook! _PLEASE!!!_

### Learning Goals
* **[G0-4]** Working with **Jupyter notebooks**.  
* **[G0-5]** Learn the basics of **programming in Python**. 
* **[G0-6]** Understand that **arrays** are multi-dimensional data structures of fixed dimensions. _Think of arrays as vectors_.  
#### Acquire Python skills
* use the `numpy` package

## 1. What is Data Science?

In the first lecture, we explored and discussed what data science is...

<!-- BEGIN QUESTION -->

### Problem 1.1

**Write-up!** What is data science? Define data science in your _own_ words. 
> **Hint**: This is defined in the very first section in the zyBook (Histroical overview).

> **Important**: In case you needed help with your answer, cite any resources you used to come up with your answer.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Data Science Use Cases
Data science is about using data to **find solutions** for a particular problem. Consider the following data science use cases: 

* automated image captions
* crime: predicting locations
* crop-yield improvment
* diagnosing heart disease
* emotion detection
* fitness tracking
* humane genome sequencing
* insurance pricing
* language translation
* personalized medicine
* smart homes
* ride sharing
* social analytics
* terrorist attack prevention


### Problem 1.2

**Write-up!** From the use cases above, pick one of them and describe
* the use case, 
* the data that is involved, and 
* how a data scientist might approach and solve the problem. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 2. Some Python Practice

For some coding problems in this class, we use a tool called otter-grader. This automatically checks the correctness of some elements of your code, and shows you some test cases you can use to assist in your development of a solution. 

Please note that this **is not** all encompassing and that many questions will **also** be graded for style, or have hidden tests that you cannot see! These questions will be marked with a note saying "*Note: This question is graded on more than public tests alone*". 

For example, the public test for the next question checks if your new string is shorter than the original, and the hidden test checks if your answer matches the intended solution. These public tests tell you what they are looking for if you fail them, so they're used to **guide** you to the solution.

Please note that the variable names we give you are what we are testing for, so **do not change these!**

### Problem 2.1 

**Do this!** Get all but the first four characters of the following string without counting its length: \
*Note: This question is graded on more than public tests alone.*

```python
my_string = "This is a really long string and you have now clue on how long it is!"

shorter_string = ...
```

```python
grader.check("q2a")
```

### Problem 2.2

Get the second to fourth elements of `a_list` (reminder: the first element of a list is not index 1, etc.).

```python
a_list = [1, 2, 3, 4, 5, 6]

new_list = ...
```

```python
grader.check("q2b")
```

### Problem 2.3

In this problem we will get some more practice using the most common syntax constructs in Python, specifically _control flow_ and _for loops_. To do this, we will tackle the notorious "FizzBuzz" question that pops up on coding assessments sometimes.

**Do this!** Complete the following function so that it returns a sequence of numbers 1 through n, where `n` is an argument. However, for each multiple of 3, add "Fizz" instead of the number, for each multiple of 5, add "Buzz" instead of the number, and for numbers which are multiples of both 3 and 5, add "FizzBuzz" instead of the number. We will check your solution for you as you go — try "running" (▶️ or `shift + enter`) the cell below and take a look at the output.

*Note: This question is graded on more than public tests alone*

> **Hint**: Start by creating a for loop to add every even number in the interval `[1:n]` to a list; you'll need to create an appropriate `iterable` first.

```python
def fizzbuzz(n):
    '''
    Returns a sequence of numbers 1 through, and including, N, except when the number
    is a multiple of 3, 5 or of both. In these cases, the sequence entry shoudl be "Fizz",
    "Buzz", and "FizzBuzz", respectively. Return the sequence as a list. 
    
    >>> fizzbuzz(5)
    
    [1, 2, 'Fizz', 4, 'Buzz']
    '''
    
    sequence = []
    ...
    return sequence

fizzbuzz(15)
```

```python
grader.check("q2c")
```

### Problem 2.4

A _hailstone sequence_ is a fun little sequence that starts with a positive _integer_ $n$. Then,
1. if $n$ is even, divide it by 2
2. if $n$ is odd, multiply it by 3 and add 1
3. continue this process until $n$ is 1.

**Do this!** Complete the following function so that it creates a list of the hailstone sequence starting at integer `n`. It should also include `n` itself. _**Hint**: use floor division `//` to get integers when dividing. Also, if you run into infinite loops, hit `Kernel > Interrupt Kernel` in the top bar._

```python
def hailstone(n):
    '''
    Returns a list of the hailstone sequence starting with n.
    
    >>> hailstone(10)
    [10, 5, 16, 8, 4, 2, 1]
    '''
    
    sequence = []
    ...
    return sequence
    
hailstone(11)
```

```python
grader.check("q2d")
```

### Problem 2.5

**Do this!** Use an appropriate function with a lambda function as its input to find the student with the soonest graduation date. Assign the result to the variable `soonest_grad_date`.

```python
roster = [
    ['Billy',  21, 2024],
    ['Meghan', 18, 2023],
    ['Jeff',   21, 2021],
    ['Alex',   50, 2024],
    ['Cate',   21, 2023]
]

soonest_grad_date = ...
```

```python
grader.check("q2e")
```

## 3. Timing Comparison

There were some questions about why we emphasize NumPy functions and indexing so intensely. To help explain and motivate this, let's do some performance comparisons.

```python
# run me!
import numpy as np
```

For this experiment, we will be testing various operations with random arrays. Here is an example of how you can generate a random array using `np.random.rand`.

```python
np.random.rand(10)
```

### Problem 3.1

**Do this!** Complete the function below so that it returns a random array of size `n`. Assign your new array to the `result` variable.

```python
def generate_random_array(n):
    '''Returns a random array of a given shape.'''
    
    result = ...
    
    return result
```

```python
grader.check("q3a")
```

Let's try it out! You should see an array similar to the one from earlier.

```python
generate_random_array(10)
```

### 3.1 Summing Arrays

Let's try a sum operation. There are two ways of computing the sum of an array using built-in features. The first is to use a `for` loop and iterate through the values in an array. The second is to use the `sum` function.

### Problem 3.2

**Implement this!** Complete the function below so that it returns the sum of an array `a`, making sure to use a `for` loop. Again, assign your result to the `result` variable.

```python
def loop_sum(a):
    '''Computes the sum of array A using a for loop'''
    
    ...
    
    return result
```

```python
grader.check("q3b")
```

<!-- BEGIN QUESTION -->

### Problem 3.3

Now that we have a working `for` loop sum implementation, `loop_sum`, let's compare its performance to both the Python built-in `sum` function and the NumPy `np.sum` function. 
> We will use an [IPython magic command 🧙‍♀️](https://ipython.readthedocs.io/en/stable/interactive/tutorial.html#magic-functions) (I'm not kidding) called `%timeit`. This function will run a given command repeatedly and report back the mean runtime and standard deviation.

```python
another_array = generate_random_array(10000)

%timeit loop_sum(another_array)
%timeit sum(another_array)
%timeit np.sum(another_array)
```

**Write up!** Let's study this output. Answer the following questions: 
1. Which way of implementing the sum is most efficient?
2. How much faster is the most efficient way of computing the sum compared to the other implementations?  
3. What would happen if you added more dimensions to your array (think _analyzing a much bigger dataset_)? 
4. What do these results tell us in terms of implementing solutions to data sceince problems?
5. Assume that the fastest way of implemeting the sum would take 1 minute. How much time would it take to compute the result using the other implementations? Report your result in hours. Report the runtimes you base your computation off, since those numbers will change when you rerun the cell above (which is likley unavoidable). 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### 3.2 Finding a Value

In Problem 2 we needed to find the age and class of the student from a roster who would graduate first. Let's use this set up to do another comparison. Here is the data that we worked with.

```python
names = ['Billy','Meghan','Jeff','Alex','Cate']
roster = np.array([
    [21, 2021],
    [18, 2020],
    [21, 2019],
    [50, 2021],
    [21, 2020]
])
```

Above we used the Python built-in `min` function with a lambda to accomplish this for a `List` version of the roster. What would happen if we had used the same method to do the same for a NumPy array?

```python
min(roster, key=lambda student: student[1])
```

An appropriate NumPy equivalent of the code is this:

```python
roster[np.argmin(roster[:, 1])]
```

We can find the student's name by doing the following:

```python
names[np.argmin(roster[:, 1])]
```

<!-- BEGIN QUESTION -->

### Problem 3.4

Before we check out the performance of each of these implementations, let's expand our roster a bit. In the following cell, we generate a new roster with entries for 1000 students.

```python
roster = np.array([np.random.randint(16, 100, size=1000),
                   np.random.randint(2018, 2022, size=1000)]).T
```

Here's a preview of the new roster containing the first ten rows in the array:

```python
roster[:10, :]
```

Now we are ready to evaluate these implementations using the `%timeit` magic command from earlier. 

```python
%timeit min(roster, key=lambda student: student[1])
%timeit roster[np.argmin(roster[:, 1])]
```

**Write up!** Let's study this output by answering the following questions: 
1. Which way of implementing the information retrieval is most efficient?
2. How much faster is the more efficient computation compared to the slower implementation?  
3. What would happen if you added more dimensions to your array (think _analyzing a much bigger dataset_)? 
4. What do these results tell us in terms of using Python versus NumPy bulit-in functions in solutions to data sceince problems?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 4. Images are Arrays!

AKA, more [NumPy](https://numpy.org) practice — so exciting, right? In this problem we will expand on the generality of NumPy's arrays; specifically, we will be looking at how we can think of images as arrays of pixel values.

But first, let's just do a quick review of some basic operations with NumPy.

### Problem 4.1

**Do this!** Create a three dimensional array of ones with shape $3 \times 3 \times 2$ and assign it to a variable called `ones_array`. _**Hint**:_ you can view the resulting array by evaluating the variable (write the name on its own line at the end of the cell; remember that only the value of the last expression in the cell will be displayed).

```python
import numpy as np

ones_array = ...
```

```python
grader.check("q4a")
```

### Image Time!

![image](utility/data/example-image.png)

This will be the image we are working with in the next probelms.

Without getting too technical, digital images are composed of pixels arranged in a rectangular grid. Each pixel has an intensity value within some range (between 0 and 1 in this case) that specifies how dark or light it should be. Shades of gray (in a black and white image) can be specified by a single value whereas colors are typically specified by a tuple of three values (see the [RGB color model](https://en.wikipedia.org/wiki/RGB_color_model)). Both types of images can be represented as multidimensional arrays.

_Note: image was taken by Erik Mclean and sourced from [Unsplash](https://unsplash.com/photos/BIheydot8Jw)._

Let's try loading the image into an array. We've hidden how this is done, but feel free to check out `utility/util.py` if you are curious.

```python
from utility.util import load_image, show_array_as_image

image = load_image('utility/data/example-image.png')
image
```

```python
show_array_as_image(image)
```

### Problem 4.2

**Do this!** Retrieve the shape of the array. 

*Note: This question is graded on more than public tests alone*

_**Hint**: the property you are looking for has the same name as what you are looking for._

```python
shape = ...
```

```python
grader.check("q4bi")
```

<!-- BEGIN QUESTION -->

**Write-up!** How do the dimensions of the array relate to the axes (height and width) of image? How were you able to tell?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Problem 4.3

**Do this!** Compute the average intensity of each "row" of pixels in the image and store the values in `average_row_intensities`. Do the same for each "column" of pixels in the image and store the values in `average_column_intensities`. You should only use NumPy's vectorized operations. No loops!

*Note: This question is graded on more than public tests alone*

_**Hint**: Which axis do you want to perform the operation across? How many values do you expect to get from each operation? Are the shapes of the resulting arrays consistent with your expectations?_

```python
average_row_intensities = ...
average_column_intensities = ...
```

```python
grader.check("q4ci")
```

<!-- BEGIN QUESTION -->

Let's try to get a better feel for what is going on. How would we find the darkest pixel in the image?

**Do this!** Get the indicies of the row and column with smallest average value. Print these values (row first). _**Hint**: refer back to Lab 0._

```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Write-up!** Are these values consistent with the image? What is located at the pixel with these coordinates? Here is the image again.

```python
show_array_as_image(image)
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Problem 4.4

Finally, let's try adjusting the brightness or [exposure](https://en.wikipedia.org/wiki/Exposure_(photography)) of the image. Remember that the `image` is represented as an array containing pixel intensities, where darker colors have values smaller than lighter colors. Thus, we can make an image "brighter" by making the pixel intensity values larger.

**Do this!** Make the image **half as bright** and store the new image in `darker_image`.

```python
darker_image = ...
```

```python
grader.check("q4d")
```

```python
# now let's look at our new image
show_array_as_image(darker_image)
```

That's it! Follow the submission instructions below:
* Restart your kernel and run all the cells
* Save your notebook
* Upload the .ipynb notebook directly to Gradescope

```python

```

Tests for q2a (points: 2)
```
>>> len(shorter_string) == len(my_string) - 4
True
```


Tests for q2b (points: 2)
```
>>> new_list == [2, 3, 4]
True
```


Tests for q2c (points: 10)
```
>>> fizzbuzz(8) == [1, 2, 'Fizz', 4, 'Buzz', 'Fizz', 7, 8]
True
```
```
>>> fizzbuzz(3) == [1, 2, 'Fizz']
True
```
```
>>> fizzbuzz(15)[14] == 'FizzBuzz'
True
```


Tests for q2d (points: 10)
```
>>> hailstone(11) == [11, 34, 17, 52, 26, 13, 40, 20, 10, 5, 16, 8, 4, 2, 1]
True
```
```
>>> hailstone(10) == [10, 5, 16, 8, 4, 2, 1]
True
```


Tests for q2e (points: 5)
```
>>> isinstance(soonest_grad_date[1], int)
True
```
```
>>> isinstance(soonest_grad_date[0], str)
True
```
```
>>> soonest_grad_date[2] == 2021
True
```
```
>>> soonest_grad_date[1] == 21
True
```
```
>>> soonest_grad_date[0] == 'Jeff'
True
```


Tests for q3a (points: 3)
```
>>> generate_random_array(5).size == 5
True
```
```
>>> isinstance(generate_random_array(20), np.ndarray)
True
```


Tests for q3b (points: 5)
```
>>> an_array = generate_random_array(100)
>>> loop_sum(an_array) == sum(an_array)
np.True_
```


Tests for q4a (points: 3)
```
>>> ones_array.shape == (3, 3, 2)
True
```
```
>>> np.sum(ones_array) == ones_array.size
np.True_
```


Tests for q4bi (points: 1)
```
>>> shape == (250, 374)
True
```


Tests for q4ci (points: 4)
```
>>> np.isclose(average_row_intensities.sum(), 205.15002) == True
np.True_
```
```
>>> np.isclose(average_column_intensities.sum(), 306.90482) == True
np.True_
```


Tests for q4d (points: 3)
```
>>> np.all(darker_image < image)
np.True_
```
